"""
match guards add an if condition to a case clause:
    case n if n < 0: return "negative"
The guard is checked after the pattern matches; if it fails, the next case is tried.

Implement classify(n) using match/case with guards so it returns "negative" when n < 0, "zero" when n == 0, "small positive" when 0 < n < 10, and "large positive" otherwise.
"""

# I AM NOT DONE


def classify(n):
    return "unknown"
