"""
match guards add an if condition to a case clause:
    case n if n < 0: return "negative"
The guard is checked after the pattern matches; if it fails, the next case is tried.

Implement classify() using match/case with guards for negative, zero, small, and large.
"""

# I AM NOT DONE


def classify(n):
    return "unknown"
