"""
__repr__ should return an unambiguous developer-facing string (e.g. ClassName(args)).
__str__ returns a user-friendly display string used by print() and f-strings.
The REPL and repr() use __repr__; str() and f-strings use __str__.

Implement both methods on Point so Point(1, 2) gives the correct repr and str.
"""

# I AM NOT DONE


class Point:
    def __init__(self, x, y):
        self.x = x
        self.y = y

    def __repr__(self):
        return ""

    def __str__(self):
        return ""
