"""
List comprehensions: [expression for item in iterable if condition]
They are concise and often faster than equivalent for-loops.
Read them left to right: "give me n**2 for each n in numbers where n % 2 == 0".

Fix the comprehension to square only even numbers from numbers.
"""

# I AM NOT DONE


numbers = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
squares_of_evens = [n for n in numbers]
