"""
In Python, strings and tuples are immutable — you cannot change their contents
after creation. Lists, dicts, and sets are mutable — you can modify them in place.

Remove the line that tries to mutate an immutable object, then remove # I AM NOT DONE.
"""

# I AM NOT DONE


hello_immutable_string = "Hello"
hello_mutable_list = ["H", "e", "l", "l", "o"]

hello_immutable_string[0] = "H"
hello_mutable_list[0] = "I WORK"
