"""
Python lets you assign a single value to multiple variables at once using
chained assignment (x = y = z = value). You can also unpack a tuple or
sequence into separate variables in one statement (a, b, c = 1, 2, 3).

The stubs below are already correct — remove # I AM NOT DONE to pass.
"""

# I AM NOT DONE


x = y = z = "Hello!"
a, b, c = "a", "b", "c"
