"""
The + operator on two strings concatenates them. On two integers it adds them.
Python does not automatically coerce between incompatible types — adding an
int to a str raises a TypeError. You must convert explicitly with int() or str().

Fix one_plus_two so it evaluates to 3 (the integer sum of 1 and 2).
"""

# I AM NOT DONE


hello = "Hello "
world = "World!"
hello_world = hello + world

one = 1
two = "2"
one_plus_two = one + two
