"""
The type of a variable in Python is determined by the literal you assign.
Integers use no decimal point (1), floats include one (1.0), and bytes
literals are prefixed with b (b"Hello").

Fix the assignments: x must be an int, y must be a float, z must be bytes.
"""

# I AM NOT DONE


x = 1.0
y = 1
z = "Hello"
