"""
The type of a variable in Python is determined by the literal you assign.
Integers use no decimal point (1), floats include one (1.0), and bytes
literals are prefixed with b (b"Hello").

Fix the assignments so each variable matches the expected type.
"""

# I AM NOT DONE


x = 1.0
y = 1
z = "Hello"
