"""
Python is a dynamically-typed language — the interpreter infers each variable's
type from the value you assign. Different types have different behaviours:
int and float are numbers, str is text, bytes is raw data, bool is True/False,
and tuple is an ordered immutable sequence.

Assign each variable the value shown, using the correct Python literal for its type.
"""

# I AM NOT DONE


a_integer = 0
a_float = 0.0
a_string = ""
a_bytes = b""
a_boolean = False
a_tuple = ()
