"""
Python is a dynamically-typed language — the interpreter infers each variable's
type from the value you assign. Different types have different behaviours:
int and float are numbers, str is text, bytes is raw data, bool is True/False,
and tuple is an ordered immutable sequence.

Assign a literal of the correct type to each variable: a_integer must be an int, a_float a float, a_string a str, a_bytes a bytes, a_boolean a bool, and a_tuple a tuple. Only the type matters; the exact value is up to you.
"""

# I AM NOT DONE


a_integer = 0
a_float = 0.0
a_string = ""
a_bytes = b""
a_boolean = False
a_tuple = ()
