# I AM NOT DONE

from functools import reduce

numbers = [1, 2, 3, 4, 5]

# Fix: use reduce to compute the product of all numbers
product = 0  # Fix this

# Fix: use reduce to find the maximum value (without using the built-in max())
maximum = numbers[0]  # Fix this

# DON'T EDIT THE TESTS
assert product == 120
assert maximum == 5
