"""
map(func, iterable) applies func to every item and returns a lazy iterator.
filter(pred, iterable) keeps items where pred returns True.
Both return iterators — wrap with list() to get a concrete list.

Fix squares, evens, and even_squares using map() and filter().
"""

# I AM NOT DONE


numbers = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]

squares = list(numbers)
evens = list(numbers)
even_squares = list(numbers)
