"""
Dicts map keys to values. Use dict[key] = value to add or update, del dict[key]
to remove a key, and dict.get(key, default) to safely read — returning a default
instead of raising KeyError when the key is absent.

Update, remove, and safely read from the prices dict.
"""

# I AM NOT DONE


prices = {"apple": 0.99, "banana": 0.59, "cherry": 2.99}
mango_price = 0.0
