"""
Lists are ordered, mutable sequences. You can change elements by index,
add items with append(), and remove the last item with pop().

Start with the list [0, "hello", True, ["other list"]], then apply operations
so the final list has length 4 with specific values at each position.
"""

# I AM NOT DONE


my_list = [0, "hello", True, ["other list"]]
