"""
Lists are ordered, mutable sequences. You can change elements by index,
add items with append(), and remove the last item with pop().

Starting from my_list = [0, "hello", True, ["other list"]], reassign index 0 to 1, index 1 to "world", and index 2 to the boolean False, using my_list[idx] = value. Keep the list at length 4 — do not append, pop, or remove any element.
"""

# I AM NOT DONE


my_list = [0, "hello", True, ["other list"]]
