"""
Starred unpacking lets you capture "the rest" of a sequence into a list.
Write: first, *rest = items — first gets items[0], rest gets the remaining elements.
The star can appear anywhere: a, *mid, z = items captures head, middle, and tail.

Use starred unpacking to split numbers into first/rest and first2/middle/last.
"""

# I AM NOT DONE


numbers = [1, 2, 3, 4, 5]

first = None
rest = None

first2 = None
middle = None
last = None
