"""
Sets are unordered collections of unique values. Set operations mirror math:
- difference (a - b): elements in a but not b
- intersection (a & b): elements in both
- union (a | b): all elements from either

Use these operators to compute evens_only, both, and all_nums from two given sets.
"""

# I AM NOT DONE


evens = {2, 4, 6, 8, 10}
primes = {2, 3, 5, 7, 11}

evens_only = set()
both = set()
all_nums = set()
