# I AM NOT DONE

evens = {2, 4, 6, 8, 10}
primes = {2, 3, 5, 7, 11}

# Numbers that are in evens but NOT in primes
evens_only = set()  # Fix this

# Numbers that appear in BOTH sets
both = set()  # Fix this

# All unique numbers from either set
all_nums = set()  # Fix this

# DON'T EDIT THE TESTS
assert evens_only == {4, 6, 8, 10}
assert both == {2}
assert all_nums == {2, 3, 4, 5, 6, 7, 8, 10, 11}
