"""
Type annotations document what types a function expects and returns.
Syntax: def func(param: int, other: str) -> bool:
They don't enforce types at runtime, but help IDEs and type-checkers catch errors early.

Add type annotations to add() (int params, int return) and greet() (str param, str return).
"""

# I AM NOT DONE


def add(a, b):
    return a + b


def greet(name):
    return f"Hello, {name}"
