"""
NamedTuple combines named-attribute access with tuple positional access.
class Point(NamedTuple): x: float; y: float
Instances are immutable (like tuples) and support both p.x and p[0].

Add x: float and y: float fields to Point.
"""

# I AM NOT DONE

from typing import NamedTuple


class Point(NamedTuple):
    pass
