"""
*args collects all extra positional arguments into a tuple.
**kwargs collects all extra keyword arguments into a dict.
Use them when you don't know in advance how many arguments a caller will pass.

Implement total() to sum all positional args and describe() to return all kwargs as a dict.
"""

# I AM NOT DONE


def total(*args):
    return 0


def describe(**kwargs):
    return {}
