# I AM NOT DONE

def total(*args):
    # Return the sum of all positional arguments
    return 0  # Fix this


def describe(**kwargs):
    # Return a dict containing all keyword arguments
    return {}  # Fix this


# DON'T EDIT THE TESTS
assert total(1, 2, 3) == 6
assert total(10, 20) == 30
assert total() == 0
assert describe(name="Ada", lang="Python") == {"name": "Ada", "lang": "Python"}
