"""
lambda creates a small anonymous function inline: lambda x: x * 2.
Use lambdas for short one-expression callbacks like sort keys or map/filter arguments.
For anything longer or reusable, prefer a named def.

Fix the three lambdas so they compute the correct expression.
"""

# I AM NOT DONE


square = lambda x: x

is_even = lambda x: True

second = lambda t: t
