"""
lambda creates a small anonymous function inline: lambda x: x * 2.
Use lambdas for short one-expression callbacks like sort keys or map/filter arguments.
For anything longer or reusable, prefer a named def.

Fix the three lambdas: square must return x squared (x ** 2), is_even must return whether x is even as a bool (x % 2 == 0), and second must return the second element of the tuple t (t[1]) so it works as a sort key.
"""

# I AM NOT DONE


square = lambda x: x

is_even = lambda x: True

second = lambda t: t
