"""
A generator function uses yield instead of return. Calling it returns a lazy
generator object; values are produced one at a time when iterated.
Python pauses at each yield and resumes when the next value is requested.

Make countdown() a generator that yields n, n-1, ..., 1 (stopping at 1).
"""

# I AM NOT DONE


def countdown(n):
    return []
