"""
Generators can be infinite — just never stop yielding.
Use itertools.islice(gen, n) to take only the first n values without consuming the rest.
Infinite generators are ideal for sequences like natural numbers or event streams.

Make naturals() an infinite generator that yields start, start+1, start+2, ...
"""

# I AM NOT DONE

from itertools import islice


def naturals(start=1):
    return
