# Public functions should have a docstring describing what they do.
# Ruff's pydocstyle rule D103 flags any public function with no docstring.
#
# Add a one-line docstring to `add` as its first statement.
# The docstring must end with a period (e.g. """Return the sum of a and b.""").

# I AM NOT DONE


def add(a, b):
    return a + b
