# A one-line docstring should be tidy. Ruff's pydocstyle rules check that it:
#   D200 - fits on one line (opening/closing quotes on the same line),
#   D210 - has no surrounding whitespace inside the quotes,
#   D400 - ends with a period.
#
# Fix the docstring below so it reads: """Return the square of n."""

# I AM NOT DONE


def square(n):
    """ square n """
    return n * n
