# A module (file) should start with a docstring describing it.
# Ruff's pydocstyle rule D100 flags a module with no docstring.
# These leading "#" lines are comments, not statements, so they do NOT
# count as a module docstring.
#
# Add a module docstring as the FIRST statement of the module (a triple-quoted
# string placed right after these comments, before `GREETING`). It must end
# with a period.

# I AM NOT DONE

GREETING = "hello"
